---
id: 5d8a4cfbe6b6180ed9a1c9de
title: Step 1
challengeType: 0
dashedName: step-1
---

# --description--

Ласкаво просимо до проекту інформаційної панелі(dashboard)! Ви будете використовувати JavaScript бібліотеку візуалізації даних, D3, щоб створити візуалізацію ваших підписників у соціальних мережах. Він складатиметься з лінійного графіка, секторної діаграми та легенди.

Спочатку вам потрібно створити файл HTML. Почніть із додавання <!DOCTYPE html> у верхній частині файлу, щоб повідомити браузеру, який тип документа він читає.

# --hints--

test-text

```js
assert(/<!DOCTYPE\s+html\s*>/gi.test(code));
```

# --seed--

## --seed-contents--

```html
```

# --solutions--

```html
<!DOCTYPE html>
```
